# Lab 02: CSS 1

Welcome to the second lab of **CSS 1**!

In this self-paced lab, you'll get hands-on practice with the following **concepts**:

- Defining and using variables.  
- Implementing and evaluating **conditional** statements.  
- Implementing and evaluating **loops**.

You'll also get some practice **debugging** code.

### Note on hidden tests

Some of the questions in this lab will have **hidden tests**. This means that in addition to the `assert` statements you see in the code, there *hidden* tests. The point of these is to check whether your code **generalizes** to other cases.

The lab will notify you when/where there is a hidden test, even if you can't see it.


### Guide to CSS 1 Labs

Programming can be challenging, whether you're experienced or completely new. But it can also be a lot of fun. If you find yourself overwhelmed or frustrated with a question, remember: we're here to help!

- If you're working on this in section, feel free to ask your TA for guidance.  
- Labs are also **collaborative**, so you can also ask your fellow classmates. In fact, we encourage it! 
- If you're stuck, you can also post questions on Piazza.

## Part 1: Variables

> A **variable** stores a particular value, like a "container".

Variables can be **declared** using the assignment operator (`=`). Once a variable has been declared, you can reference it using the variable name.

### Declaring variables

In this section, you'll get some more practice declaring variables of different types.

In the cell below, declare the following variables:

- `my_first_bool`: set it equal to either `True` or `False`.  
- `my_number`: set it to an `int`, like `2` or `200`. 
- `my_word`: set it to the `str` `"css"`. 
- `my_none`: set it to `None`. 

<!-- grade_id: cell-f3ef5e43fd512d88, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION
```

<!-- grade_id: cell-72872bb8d1b558db, points: 0.5, grade: true, solution: false -->
```python
assert my_first_bool == True or my_first_bool == False
assert type(my_first_bool) is bool
```

<!-- grade_id: cell-7bccbd3c9853ec6d, points: 0.5, grade: true, solution: false -->
```python
assert type(my_number) is int
```

<!-- grade_id: cell-ddde4e69be5e663b, points: 0.5, grade: true, solution: false -->
```python
assert my_word == "css"
assert type(my_word) is str
```

<!-- grade_id: cell-2be573e002c4a85b, points: 0.5, grade: true, solution: false -->
```python
assert my_none is None
```

### Operating with variables

In this section, you'll get hands-on practice **using** variables in some basic operations.

The cell below has defined two variables, `x1` and `x2`. 

Create a new variable called `x3`, which is equal to the *sum* of these variables. 

```python
x1 = 5
x2 = 2
```

<!-- grade_id: cell-21b2aa4fc6110b5a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION
```

The cell below has at least one hidden test.

<!-- grade_id: cell-17c865a083c5c38e, points: 1, grade: true, solution: false -->
```python
## Includes Hidden Test
assert type(x3) is int
### BEGIN HIDDEN TESTS
assert x3 == 7
### END HIDDEN TESTS
```

Create a new variable called `x5`, which is equal to the *product* of `x1` and `x2`.

<!-- grade_id: cell-68c57ef6a6521a9e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION
```

<!-- grade_id: cell-38f2f04dcaa41bee, points: 1, grade: true, solution: false -->
```python
assert x5 == 10
```

The cell below has another new variable, called `x4`.

Now create another new variable called `comp1`. Set this variable equal to a `bool` value reflecting whether `x4` is larger than `x3`. 

**Hint**: Use the `>` operator; the value should be `True` if `x4` is larger, and `False` if `x4` is smaller.

```python
x4 = 100
```

<!-- grade_id: cell-aafd6e67b8054ae9, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
comp1 = x4 > x3
### END SOLUTION
```

The cell below has at least one hidden test.

<!-- grade_id: cell-e79121830ea40e91, points: 1, grade: true, solution: false -->
```python
## Includes Hidden Test
assert type(comp1) is bool
### BEGIN HIDDEN TESTS
assert comp1 == True
### END HIDDEN TESTS
```

Now let's try a more complex comparison. Is `x4` larger than the product of `x1`, `x2`, and `x3`? Set the result of this comparison to `comp2`. 

**Hint**: Use the `>` operator; the value should be `True` if `x4` is larger, and `False` if `x4` is smaller.

<!-- grade_id: cell-43c9b874561e1ab8, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION
```

<!-- grade_id: cell-68a6aa62f170d67d, points: 1, grade: true, solution: false -->
```python
assert comp2 == True
```

#### The `%` operator

The `%` operator returns the **remainder** of dividing one number by another. For example, `3 % 2` would return `1`.

In the cell below, create a variable called `remainder`, which is equal to the remainder of dividing `100` by `12`.

<!-- grade_id: cell-47b88dc2404ecfe1, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION
```

<!-- grade_id: cell-bbe6c9f70a0f9e0c, points: 1, grade: true, solution: false -->
```python
assert type(remainder) is int
assert remainder == 4
```

## Part 2: Conditional Statements

A [conditional statement](https://ucsd-css1-introduction.github.io/overview/lectures/05-conditionals.html) is a statement that checks whether some **condition** is `True` or `False`.

This allows you to execute a piece of code *depending* on the value of that condition.

### Modifying a conditional statement

In the cells below, there is a **conditional statement**, which makes various changes to a variable called `x`, depending on the value of `c1` and `c2`.

In the cell immediately below, create two **boolean** variables: `c1` and `c2`. Find the correct value (`True` vs. `False`) for each variable in turn, such that when the conditional statement is run, `x == 30`.

Note that the code should look something like:

```python
c1 = bool_value_here
c2 = bool_value_here
```

But with the `bool_value_here` set to either `True` or `False`.

<!-- grade_id: cell-c85dfe407c1a73be, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION
```

```python
x = 10
if c1 and c2:
    x *= 2
elif c1 and not c2:
    x *= 3
elif not c1 and c2:
    x *= 4
else:
    x *= 5
```

<!-- grade_id: cell-22bc884050d41370, points: 1, grade: true, solution: false -->
```python
assert type(c1) is bool
assert type(c2) is bool
assert x == 30
```

### Evaluating a conditional statement

It's also important to get practice learning how to **evaluate** the outcome of running a conditional statement. 

In the cells below, there is another conditional statement, which will result in a particular value being assigned to the variable `result`.

To test whether you can tell what value will be assigned to `result`, create a variable called `predicted_outcome`. The auto-grader will then test whether `predicted_outcome == result`.

<!-- grade_id: cell-558a47713e782a17, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION
```

```python
if 100 % 3 == 0 and 100 % 5 == 0:
    result = "fizzbuzz"
elif 100 % 5 == 0:
    result = "buzz"
elif 100 % 3 == 0:
    result = "fizz"
else:
    result = "nothing"
```

<!-- grade_id: cell-f16e09c2e14bb2c4, points: 1, grade: true, solution: false -->
```python
assert predicted_outcome == result
```

## Part 3: Loops

> A [loop](https://ucsd-css1-introduction.github.io/overview/lectures/06-loops.html) is a way to repeat the same piece of code many times.

In this section, you'll answer two questions that involve **evaluating** and **assessing** what a loop will do.

### What will a loop return?

The code in the cells below runs a `for` loop over a `list`. The result of running this code will return a certain value for `result_loop`.

To test whether you can tell what value will be assigned to `result_loop`, create a variable called `predicted_outcome_loop`. The auto-grader will then test whether `predicted_outcome_loop == result_loop`.

<!-- grade_id: cell-517995be41f14eb0, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION
```

```python
result_loop = 0
for num in [1, 3, 6, 8, 11, 15]:
    if num % 2 == 0:
        result_loop += 1
result_loop
```

Output:
```
2
```

<!-- grade_id: cell-d51e513c0bf42398, points: 1, grade: true, solution: false -->
```python
assert predicted_outcome_loop == result_loop
```

## Part 4: Submit!

Once you've completed all the cells above (saving regularly):

- Click "Validate". This will run a check to determine whether you've passed all visible tests. 
- Once you've validated the assignment, you should now have an option to "Submit" the assignment (next to where the assignment is stored in your directory). Click this.
- This will now show up under your "Submitted Assignments" section.

If you have any trouble accessing or submitting the assignment, please check in with your TA!