# Problem Set 1

This is your first **problem set** for CSS 1.

It will cover the topics discussed in the first two weeks of class, including:

- Working with different **variable types**.  
- Basic **arithmetic operations**.  
- Basic **control flow**, such as *conditional statements* and *loops*.  

## Things to keep in mind

### Independent work

Note that unlike labs, problem sets should be **completed independently**.

### Note on hidden tests

Your problem set will be graded (mostly) with an **auto-grader**. Visible tests can be seen directly in the Jupyter notebook, and can be identified by their `assert` statements.

Additionally, some of the questions in this problem set will have **hidden tests**. This means that in addition to the `assert` statements you see in the code, there are *hidden* tests. The point of these is to check whether your code **generalizes** to other cases. The problem set will notify you when/where there is a hidden test, even if you can't see it.


### Q1. Creating variables (6 pts. total)

In this first part of the problem set, your goal is to **create variables** of different types, with different values.

In the cell below, create the following variables (make sure to use the right *name* for each variable):

- A variable called `my_int` that stores the number `42`.  
- A variable called `my_int2` that stores the number `45`.
- A variable called `my_str` that stores the string `"Problem Set 1"`.  
- A variable called `my_bool` that stores the boolean value `True`.  
- A variable called `my_other_bool` that stores the (boolean) result of checking whether `my_int > my_int2`.  
- A variable called `my_nothing` that stores `None`.  

<!-- grade_id: cell-08726727263fcaf7, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-ed1938bc2aee1f5f, points: 1, grade: true, solution: false -->
```python
assert my_int
assert my_int == 42
```

<!-- grade_id: cell-14aa06a3b415fdcf, points: 1, grade: true, solution: false -->
```python
assert my_int2
assert my_int2 == 45
```

<!-- grade_id: cell-7a230f5f37215a31, points: 1, grade: true, solution: false -->
```python
assert my_str
assert my_str == "Problem Set 1"
```

<!-- grade_id: cell-77efd2e19f11577f, points: 1, grade: true, solution: false -->
```python
assert my_bool
assert my_bool == True
```

<!-- grade_id: cell-dd02685519a716aa, points: 1, grade: true, solution: false -->
```python
assert my_other_bool == False
```

<!-- grade_id: cell-d32b6f85c3350844, points: 1, grade: true, solution: false -->
```python
assert my_nothing is None
```

### Q2. Using variables for arithmetic (6 pts. total)

In this section, you'll create some *more* variables, but these will involve some more complicated **arithmetic operations**.

Specifically, cell below defines a set of starting variables; you'll need to use these variables in arithmetic operations to create new variables.

Using `a`, `b`, `c`, and `d`, create the following variables:

- Set `ans1` to the result of adding `a`, `b`, `c`, and `d` together.  
- Set `ans2` to the result of multiplying `a`, `b`, `c`, and `d` together.
- Set `ans3` to the result of raising `a` to the power of `d`. (Recall that `**` can be used for exponentiation.)
- Set `ans4` to the result of dividing `ans3` by `ans1`.
- Set `ans5` to the result of checking whether `ans1 > ans4`.
- Set `ans6` to the result of checking whether `ans2` is **even**. Remember that you can use the `%` operator to check whether a given number is even (i.e., whether the *remainder* of that number divided by `2` is `0`.)

```python
a = 2
b = 3
c = 5
d = 10
```

<!-- grade_id: cell-7d6f3f84f290a14e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-6b397e928cfb38a9, points: 1, grade: true, solution: false -->
```python
## Test for ans1
assert ans1
assert ans1 == 20
```

<!-- grade_id: cell-800ebc4207e3bbe0, points: 1, grade: true, solution: false -->
```python
## Test for ans2
## This cell contains a hidden test!
assert ans2
assert type(ans2) is int
### BEGIN HIDDEN TEST
assert ans2 == 300
### END HIDDEN TEST
```

<!-- grade_id: cell-e8de209a5eb620cf, points: 1, grade: true, solution: false -->
```python
## Test for ans3
## This cell contains a hidden test!
assert type(ans3) is int
### BEGIN HIDDEN TEST
assert ans3 == 1024
### END HIDDEN TEST
```

<!-- grade_id: cell-1b50c5d48d467e38, points: 1, grade: true, solution: false -->
```python
## Test for ans4
## This cell contains a hidden test!
assert ans4
### BEGIN HIDDEN TEST
assert ans4 == 51.2
### END HIDDEN TEST
```

<!-- grade_id: cell-852c227ff7112c5e, points: 1, grade: true, solution: false -->
```python
## Test for ans5
## This cell contains a hidden test!
assert type(ans5) is bool
### BEGIN HIDDEN TEST
assert ans5 == False
### END HIDDEN TEST
```

<!-- grade_id: cell-b0baf2d185af4dc6, points: 1, grade: true, solution: false -->
```python
## Test for ans6
## This cell contains a hidden test!
assert type(ans6) is bool
### BEGIN HIDDEN TEST
assert ans6 == True
### END HIDDEN TEST
```

### Q3. Conditional statements and loops (3 pts.)

The problem has three parts:

- Write a `for` loop that iterates through `range(1, 11)` and counts the number of even numbers. Store this value in `num_evens`.  
- Additionally, calculate the **sum** of all those even numbers. Store this value in `sum_evens`.  
- Finally, divide `sum_evens` by `num_evens`. Store this value in `mean_evens`. (This is the **average**.)

```python
num_evens = 0 ### Override this number in the solution below
sum_evens = 0 ### Override this number in the solution below
mean_evens = 1 ### Override this number in the solution below
```

<!-- grade_id: cell-97664b842d673cd8, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: cell-4819cbcacc84411d, points: 1, grade: true, solution: false -->
```python
assert num_evens
assert num_evens == 5
```

<!-- grade_id: cell-f110dafa41d19633, points: 1, grade: true, solution: false -->
```python
### This cell has a hidden test
assert sum_evens
### BEGIN HIDDEN TESTS
assert sum_evens == 30
### END HIDDEN TESTS
```

<!-- grade_id: cell-c56d21d89b4d70aa, points: 1, grade: true, solution: false -->
```python
### This cell has a hidden test
assert mean_evens
### BEGIN HIDDEN TESTS
assert mean_evens == 6
### END HIDDEN TESTS
```

### Submit!

Once you've completed all the cells above (saving regularly):

- Click "Validate". This will run a check to determine whether you've passed all visible tests.
- Once you've validated the assignment, you should now have an option to "Submit" the assignment (next to where the assignment is stored in your directory). Click this.
- This will now show up under your "Submitted Assignments" section.

If you have any trouble accessing or submitting the assignment, please check in with your TA!